# Podstawy Sztucznej Inteligencji 2018/2019



Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

```python

```

---

## Iloczyn skalarny

Mając wektory $n$-elementowe $\mathbf{a}$ i $\mathbf{b}$ napisać funkcję obliczającą ich iloczyn skalarny:

$$s =\mathbf{a}\cdot \mathbf{b} =  \sum_{i=1}^{n} a_i b_i $$ 

 - za pomocą iteracji po elementach
 - używając funkcjonalności `numpy`

```python
import numpy as np
```

<!-- grade_id: cell-06e5e3a8f5681458, points: 0, grade: false, solution: true -->
```python
def s1(a,b):
    '''
    Funkcja oblicająca iloczyn skalarny dwóch wektorów.
    '''
    if len(a) != len(b):
        raise ValueError
    summ = 0
    for i in range(len(a)):
        summ += a[i] * b[i]
    return summ
def s2(a,b):
    '''
    Funkcja obliczająca iloczn skalarny dwóch wektorów za pomocą funkcji wbudowanej.
    '''
    return np.dot(a,b)
```

```python
#Tests
for i in range(1000):
    a = np.random.randint(1,100, size = 10)
    b = np.random.randint(1,100, size = 10)
    if s1(a, b) != s2(a, b):
        print('blad')
```

<!-- grade_id: cell-96f566befa6dfa4f, points: 4, grade: true, solution: false -->
```python
assert s1([1,1,1,1],[1,1,1,1]) == 4 
assert s2([1,1,1,1],[1,1,1,1]) == 4
assert np.isscalar(s2([1,1,1,1],[1,1,1,1]))
assert np.isscalar(s1([1,1,1,1],[1,1,1,1]))
assert s2([7,7,7],[7,7,7]) == 147
assert s1([7,7,7],[7,7,7]) == 147

```

## Mnożenie macierzy przez wektor:

Mając wektor $n$-elementowy $\mathbf{x}$ oraz macierz $\mathbf{A}$ o wymiarach $m\times n$ napisać funkcję obliczającą ich iloczyn:

$$\mathbf{y}= \mathbf{A}\mathbf{x}.$$

Każdy element wektora $\mathbf{y}$ jest dany przez:

$$y_i = \sum_{j=1}^{n} a_{ij} x_j $$ 

 - za pomocą iteracji po elementach (podwójna pętla)
 - korzystając z faktu, że
 każdy element wektora $y_i$ jest iloczynem skalarnym $i$-tego rzędu macierzy $\mathbf A$ oraz wektora $\mathbf{x}$ (pojedyncza pętla)
 - używając funkcji:  `np.dot` lub `np.tensordot` (bez pętli)


<!-- grade_id: cell-fb522d133295441a, points: 0, grade: false, solution: true -->
```python
def y1(A,x):
    '''
    Funkcja mnożąca macierz przez wektor używająca 2 pętli.
    '''
    if len(A[0]) != len(x):
        raise ValueError    # Gdy liczba elementów wektora jest większa od ilości kolumn macierzy.
    a = []
    for i in range(len(A)):
        summ = 0
        for k in range(len(A[i])):
            summ += A[i][k] * x[k]
        a.append(summ)
    return np.array(a)

def y2(A,x):
    '''
    Funkcja mnożąca macierz przez wektor używając 2 pętli.
    '''
    if len(A[0]) != len(x):
        raise ValueError    # Gdy liczba elementów wektora jest większa od ilości kolumn macierzy.
    a = []
    for i in range(len(A)):
        a.append(s1(A[i], x))
    return np.array(a)
    
def y3(A,x):
    '''
    Funkcja mnożąca macierz przez wektor używając funkcji wbudowanej.
    '''
    return np.dot(A,x)
```

```python
#Tests
for i in range(1000):
    a = np.random.randint(5, size = (2,4)) # Random Matrix
    b = np.random.randint(1,100, size = 4) # Random wektor
    if not np.array_equal(y1(a,b), y3(a,b)) or not np.array_equal(y2(a,b), y3(a,b)):
        print('blad')
```

<!-- grade_id: cell-eff8a1c20f8fb28f, points: 6, grade: true, solution: false -->
```python
A = np.arange(9).reshape(3,3)
x = np.arange(3)
np.testing.assert_almost_equal(y1(A,x),np.array([ 5, 14, 23]))
np.testing.assert_almost_equal(y2(A,x),np.array([ 5, 14, 23]))
np.testing.assert_almost_equal(y3(A,x),np.array([ 5, 14, 23]))

```

<!-- grade_id: cell-982fed0f9cd60244, points: 1, grade: true, solution: false -->
```python
try:
    y1(np.arange(9).reshape(3,3),np.array([1,1]))
except ValueError:
    pass
else:
    raise AssertionError("Second dimension of matrix should be the same as dimension of vector!.")
```

## Mnożenie macierzy

Mamy dwie macierze $\mathbf{A}_{m\times n}$ i $\mathbf{B}_{n\times k}$. Napisać funkcję obliczającą ich iloczyn:

$$\mathbf{C}= \mathbf{A}\mathbf{B}$$.

Każdy element macierzy  $\mathbf{C}$ jest dany przez:

$$c_{ij} = \sum_{k=1}^{n} a_{ik} b_{kj} $$ 

 - za pomocą iteracji po elementach (potrójna pętla)
 - korzystając z faktu, że każdy element macierzy $c_{ij}$ jest iloczynem skalarnym $i$-tego rzędu macierzy $\mathbf A$ oraz  $j$-tego rzędu macierzy $\mathbf B$ (podwójna pętla)
 - używając funkcji:  `np.dot` lub `np.tensordot` (bez pętli)
 

<!-- grade_id: cell-bc5bc5d32ae48f5f, points: 0, grade: false, solution: true -->
```python
def C1(A,B):
    '''
    Funkcja mnożąca macierz przez macierz używająć 3 pętli.
    '''
    if len(A[0]) != len(B):
        raise ValueError    # Sprawdzenie czy macierz A ma tyle kolumn co macierz B wierszy.
    a = []
    for i in range(len(A)):
        a.append([])
        for l in range(len(B[i])):
            summ = 0 
            for k in range(len(A[i])):
                summ += (A[i][k] * B[k][l])
            a[i].append(summ)
    return np.array(a)
    
def C2(A,B):
    '''
    Funkcja mnożąca macierz przez macierz używająć 2 pętli oraz funkcji obliczającej iloczyn skalarny 2 wektorów
    zaimplementowanej wcześniej.
    '''
    a = []
    for i in range(len(A)):
        a.append([])
        for k in range(len(B[i])):
            a[i].append(s1(A[i], B[:,k]))
    return np.array(a)
def C3(A,B):
    '''
    Funkcja mnożąca macierz przez macierz używając funkcji wbudowanej.
    '''
    return np.dot(A,B)
```

```python
#Tests
for i in range(1000):
    a = np.random.randint(15, size = (7,8)) # Random Matrix
    b = np.random.randint(15, size = (8,9)) # Random Matrix
    if not np.array_equal(C1(a,b), C3(a,b)) or not np.array_equal(C2(a,b), C3(a,b)):
        print('blad')
```

<!-- grade_id: cell-fdb6f13441a02162, points: 6, grade: true, solution: false -->
```python
m = np.arange(9).reshape(3,3)
assert np.prod(C1(np.eye(3),m) == m)
assert np.prod(C1(m,m) == C3(m,m))
assert np.prod(C2(m,m) == C3(m,m))
```

<!-- grade_id: cell-f8df427efe4ce6de, points: 1, grade: true, solution: false -->
```python
try:
    C1(np.arange(9).reshape(3,3),np.ones((2,3)))
except ValueError:
    pass
else:
    raise AssertionError("Second dimension of matrix should be the same as dimension of vector!.")
```

## Ślad macierzy

Obliczyć ślad macierzy $n\times n$:

$$ Tr(A) = \sum_{i=1}^{n} a_{ii}$$

<!-- grade_id: cell-6cfbc4d0c3f5e8a4, points: 0, grade: false, solution: true -->
```python
def Tr(A):
    n = len(A)
    if n != len(A[0]):
        raise ValueError
    summ = 0
    for i in range(len(A)):
        summ += A[i][i]
    return summ
```

```python
#Tests
for i in range(10):
    a = np.random.randint(15, size = (8,8)) # Random Matrix
    if Tr(a) != np.matrix.trace(a):
        print('blad')
```

<!-- grade_id: cell-6fcc4a8c88b8f364, points: 2, grade: true, solution: false -->
```python
assert Tr(np.diag([1,2,3,4])) == 10
```

<!-- grade_id: cell-3a491937c81d5cfd, points: 1, grade: true, solution: false -->
```python
try:
    Tr(np.ones((2,3)))
except ValueError:
    pass
else:
    raise AssertionError("The matrix should be square matrix!.")
```

## Wyznacznik macierzy

Obliczyć wyznacznik macierzy korzystając z rozwinięcia [Laplace'a](https://pl.wikipedia.org/wiki/Rozwini%C4%99cie_Laplace%E2%80%99a)


```python
import numpy as np
A = np.random.randint(4,size=(4,4)).astype(np.float)
x = np.random.randint(4,size=4)
```

<!-- grade_id: cell-a6b07fe35339bafe, points: 0, grade: false, solution: true -->
```python
def Le(A):
    '''
    Funkcja wyznaczająca wyznacznik maciarzy kwadratowej A.
    '''
    if len(A) == 1:
        return A[0][0]
    else:
        suma = 0
        for i in range(len(A[0])):
            suma += ((-1)**(i)) * A[0][i] * Le(np.delete(np.delete(A,0,0), i, 1))
    return suma
```

```python
#Tests
for i in range(100):
    a = np.random.randint(15, size = (4,4)) # Random Matrix
    if Le(a) != round(np.linalg.det(a)):
        print('blad')
```

<!-- grade_id: cell-dac69a3e1b14c456, points: 10, grade: true, solution: false -->
```python
A = np.array([[ 1.,  3.,  0.,  0.],[ 3.,  2.,  3.,  2.],[ 1.,  2.,  2.,  2.],[ 1.,  0.,  0.,  3.]])
np.testing.assert_allclose( Le(A), np.linalg.det(A) )
```

```python

```